```python
# Initialize Otter
import otter
grader = otter.Notebook("lab00.ipynb")
```

---

<h2><center>NumPy and Autograder demo<br></center></h2>

---

```python
from hashutils import *
```

# Part 1) Tennis ball launcher
A tennis ball launcher launches tennis balls from a height $h=80$ cm off the ground, at an angle $\theta_0$, and intial speed of $v_0$. If the initial angle is $\theta_0=\pi/4$, find the speed $v_0$ so that the ball bounces at 10 meters from the launch point.

<img src="launcher.png" />

The formula for $v_0$ as a function of $d$, $\theta_0$, and $h$ is,

$$v_0 = \sqrt{ \frac{gd^2}{2\cos^2(\theta_0)(h+d\tan(\theta_0))} }$$

where $g=9.81$ m/s2 is the acceleration due to gravity.

```python
g = 9.81
h = 0.8
d = 10
theta0 = np.pi/4
v0 = ...
```

```python
grader.check("q1")
```

# Part 2) Cell phone tower

The file `pos_phones.txt` contains the $(x,y,z)$ coordinates of 200 cell phones. The file 
`pos_tower.txt` contains the $(x,y,z)$ coordinates of a cell phone tower. Find the position of the phone that is closest to the tower. 

### Part 2.1) Load the data

Use [`np.loadtxt`](https://numpy.org/doc/2.1/reference/generated/numpy.loadtxt.html) to load the data into NumPy arrays.

```python
pos_tower = ...
pos_phones = ...
```

```python
grader.check("q2p1")
```

### Part 2.2) Compute the distances from each of the phones to the cell tower.

```python
d = ...
```

```python
grader.check("q2p2")
```

### Part 2.3) Find the distance to the closest phone, as well as its  $(x,y,z)$ coordinates. 

```python
ind = ...                   # index of the closest phone
d_closest = ...             # distance to the closest phone
pos_closest = ...           # position of the closest phone
```

```python
grader.check("q2p3")
```

---

To double-check your work, the cell below will rerun all of the autograder tests.

```python
grader.check_all()
```

## Submission

Make sure you have run all cells in your notebook in order before running the cell below, so that all images/graphs appear in the output. The cell below will generate a zip file for you to submit. **Please save before exporting!**

Please submit the .zip file to Gradescope.

```python
# Save your notebook first, then run this cell to export your submission.
grader.export(pdf=False)
```

Tests for q1 (points: None)
```
>>> bool(v0 > 5 and v0 < 12)
True
```
```
>>> get_hash(v0, 4) == 'f22a50d170f47b09eca595e9cbce5e6a'
True
```


Tests for q2p1 (points: None)
```
>>> pos_tower.shape == (3,) and pos_phones.shape == (200, 3)
True
```
```
>>> get_hash(pos_tower, 4) == '0754256bded584cec121b93843b34cae'
True
```
```
>>> get_hash(pos_phones, 4) == '5e373f45e96f37b4dbf30513ce9635f8'
True
```


Tests for q2p2 (points: None)
```
>>> d.shape == (200,) and bool(np.isclose(d[0], 328.78, atol=0.01))
True
```
```
>>> get_hash(d, 4) == '0cabdd79004b2d945c1fbd12b7c714f8'
True
```


Tests for q2p3 (points: None)
```
>>> bool(d_closest < 180 and d_closest > 160)
True
```
```
>>> get_hash(d_closest, 4) == '433eefa36fe98ce9b67bb20b3d30e848'
True
```
```
>>> get_hash(pos_closest, 4) == 'b13f7a1de29e245f09756a2ba6b272e0'
True
```
