```python
# Initialize Otter
import otter
grader = otter.Notebook("lab6.ipynb")
```

---

<h1><center>SDSE Lab 6 <br><br> Ensemble methods and hyperparameter tuning </center></h1>

---




In this lab exercise we will apply several classification models to a problem in astrophysics. The problem is described [here](https://satyam5120.medium.com/predicting-a-pulsar-star-using-different-machine-learning-algorithms-d22ee8fc71b4) and [here](https://www.kaggle.com/datasets/colearninglounge/predicting-pulsar-starintermediate). It consists in labeling observations of space objects as either "pulsars" or "not pulsars", based on the properties of a so-called integrated profile and the DM-SNR curve. The definitions of an "integrated profile" and of the "DM-SNR curve" are not important for this activity.

The dataset has 8 feature columns:
1. Mean of the integrated profile.
2. Standard deviation of the integrated profile.
3. Excess kurtosis of the integrated profile.
4. Skewness of the integrated profile.
5. Mean of the DM-SNR curve.
6. Standard deviation of the DM-SNR curve.
7. Excess kurtosis of the DM-SNR curve.
8. Skewness of the DM-SNR curve.

Our goal is to choose a classification model from the ones covered in this class. The procedure will follow these steps:
1. Load the data.
2. Remove null values.
3. Compute the number of samples per class.
4. Extract a test dataset.
5. Build models:
    + Logistic regression
    + Random forest
    + AdaBoost
    + Gradient Boosted Trees
6. Select and evaluate a final model.

```python
import numpy as np
import pandas as pd
import matplotlib.pyplot as plt
from lab6_utils import unpack_gridsearch, plot_grid_result
from hashutils import *
rng_seed = 2434
```

# 1. Load the data

+ Load the data file `pulsar_data.csv` into a Pandas dataframe.
+ Save the column headers corresponding to feature names (all except the `target_class`) to the variable `feature_names`.

**Hints**:

+ [`pd.read_csv`](https://pandas.pydata.org/docs/reference/api/pandas.read_csv.html)
+ [`DataFrame.columns`](https://pandas.pydata.org/docs/reference/api/pandas.DataFrame.columns.html)

```python
rawdata = ...
feature_names = ...
```

```python
grader.check("q1")
```

# 2. Remove null values

Remove any feature columns of `rawdata` with more than zero null values. Store the resulting DataFrame as `cleandata`. Do not modify `rawdata`.

**Hints**: 

+ [`DataFrame.copy`](https://pandas.pydata.org/docs/reference/api/pandas.DataFrame.copy.html)
+ [`np.any`](https://numpy.org/doc/stable/reference/generated/numpy.any.html)
+ [`pd.isnull`](https://pandas.pydata.org/docs/reference/api/pandas.isnull.html)  
+ [`DataFrame.drop`](https://pandas.pydata.org/docs/reference/api/pandas.DataFrame.drop.html) 

```python
cleandata = ...
```

```python
grader.check("q2")
```

# 3. Number of samples per class
Find the number of data points in each of the two classes. Use these variable names:
+ `N0`... number of samples of class 0 (not a pulsar), 
+ `N1`... number of samples of class 1 (pulsar). 





What is the minimum accuracy expected of any non-trivial model? Save your answer to the variable `baseline_acc`.

```python
N0 = ...
N1 = ...
baseline_acc = ...
```

```python
grader.check("q3")
```

# 4. Extract the test dataset

Use Scikit-learn's `train_test_split` method to split `cleandata` into trainging and testing parts. Keep 90% for training and 10% for testing. 

**Notes**: 
+ Pass `random_state=rng_seed` to `train_test_split` so that the result is repeatable.

```python
from sklearn.model_selection import train_test_split
Xtrain, Xtest, ytrain, ytest = ...
```

```python
grader.check("q4")
```

# 5. Logistic regression 

This is a binary classification problem, so it makes sense to begin with logistic regression. 

## 5.1. Training

In a previous lab we found that logistic regression can benefit from a normalization (or scaling) of the inputs. 
Create a pipeline that includes a `StandardScaler` (name it `'scaler'`), followed by `LogisticRegression` model (name it `'model'`). 

The logistic regression model should be created with these parameters:
+ `solver='liblinear'`
+ `random_state=rng_seed`

Train the model. 

**Note**: 
+ Import Scikit-learn's `Pipeline`, `LogisticRegression`, and `StandardScaler` classes.

```python
...
logreg_pipe = Pipeline([
    ('scaler' , ... ) ,
    ('model' , ... )
])
...
```

```python
grader.check("q5p1")
```

## 5.2. Hyperparameter tuning with grid search

To solve an optimization problem by "grid search" is to exhaustively search for a solution over a grid of points covering the feasible space. The "hyperparameter tuning problem" is to select the values of the hyperparameters that maximize the performance of the model. In this section we will solve the hyperparameter tuning problem with grid search. Each setting of the hyperparameters will be evaluated using K-fold cross-validation (as opposed to a validation dataset).

Scikit-learn provides an implementation of grid search with K-fold cross-validation in the [`GridSearchCV`](https://scikit-learn.org/stable/modules/generated/sklearn.model_selection.GridSearchCV.html) class. See the documentation of this class for a detailed explanation. 

Our model will be the logistic regression pipeline from part 5.1. The hyperparameters and values to search over are the following. 

+ `penalty`: the regularization function. Try `'l1'` (LASSO) and `'l2'` (Ridge).
+ `C`: the regularization weight. This is equivalent to $1/\lambda$ in our notation. Try five logarithmically spaced values between 0.001 and 10 (i.e. $10^{-3}$, $10^{-2}$, $10^{-1}$, $1$, and $10$). 

Running grid search with Scikit-learn involves two steps: creating the `GridSearchCV` object, and then executing the algorithm with the `fit` method. Create a `GridSearchCV` with these arguments:
+ `param_grid=param_grid`. See explanation below.
+ `scoring='accuracy'`. Use accuracy as the performance metric.
+ `cv=3`. K=3 folds
+ `refit='accuracy'`. Train a final "best model" with the optimal hyperparameters and all of the training data. 

The list of candidate values for each of the hyperparameters is passed to `GridSearchCV` as a dictionary (`param_grid`) with key:value pairs parmeter name:list of parameter values. When working with a pipeline, the parameter name must be prefixed with the name of the model ('model' in our case), and two underscores (`__`). Hence the `param_grid` argument to  `GridSearchCV` should look like this:

```python
param_grid = {
    'model__penalty' : ...,
    'model__C' : ... }
```


The second step is to run the grid search algorithm passing the training data to the `fit` method on the `GridSearchCV` object. 

Implement grid search in the cell below.

**Note**: 

+ Import Scikit-learn's `GridSearchCV` class

```python
...
param_grid = ...
gs_logreg = GridSearchCV(...)
gs_logreg.fit(...)
```


`lab6_utils.py` contains useful functions for unpacking and plotting the result of the grid search. 

Below we use the `unpack_gridsearch` method to extract information from the grid search solution. This method returns a dictionary with the following entries:
+ `scoregrid`: The values of the accuracies over the hyperparameter grid. 
+ `best_params`: A dictionary with the best-case hyperparameter values.
+ `best_model`: The logistic regression pipeline corresponding to the best hyperparameters. 
+ `best_score`: The accuracy of the best model.

`plot_grid_result` plots the results of the grid search.


<!-- X, best_params, best_estimator, best_score -->

```python
logreg_result = unpack_gridsearch(gs_logreg)
plot_grid_result(logreg_result)
```

```python
grader.check("q5p2")
```

## 5.3. Package it into a function

Next we will repeat this process with several other classification models. To keep the code clean, we will first collect the steps into a single function, which we'll call `hypersolve`. This function receives as inputs:
1. The classifier object, e.g. `LogisticRegression(solver='liblinear',random_state=rng_seed)`
2. The `param_grid` dictionary that defines the search space for `GridSearchCV`.

It should:
1. Create a pipeline containing a standard scaler and the model.
2. Construct the `GridSearchCV` object as was done in part 5.2
3. Run the grid se

Run `fit` on the grid search object, using the training data
4. Run `unpack_gridsearch` the dictionary with results.
5. Plot the result with `plot_grid_result`
6. Return the `result` dictionary.


**Note**: This part has no autograder test. Failures in subsequent parts may be due to errors in `hypersolve`.

```python

def hypersolve(model,param_grid):
    # 1. Create the pipeline model
    pipe = Pipeline([
    ...
    ])

    # 2. Construct the `GridSearchCV` object as was done in part 5.2
    gs = GridSearchCV(...)
    
    # 3. Run the grid search
    gs = gs.fit(...)
    
    # 4. Run `unpack_gridsearch` to obtain the results dictionary.
    result = unpack_gridsearch(...)

    # 5. Plot the result with `plot_grid_result`
    plot_grid_result(...)

    # 6. return result
    return ...
```

```python
# test your code
...
```

# 6. Random forest

Use grid search to optimize a random forests model over the following parameter ranges:
+ `max_features`: This is the number of features to search over when splitting a node. Test 3 features and all features. 
+ `n_estimators`: Test 2, 42, and 82 trees. 

```python
from sklearn.ensemble import RandomForestClassifier

param_grid_rf = ...

model = RandomForestClassifier(random_state=rng_seed)

result_rf = hypersolve(...)
```

```python
grader.check("q6")
```

# 7. AdaBoost

Repeat the grid search, this time using an AdaBoost model. Use the following values for the parameter grid:
+ `learning_rate: [0.01,0.1]`. $\lambda$ in the notation of the lecture.
+ `n_estimators: [20,60,100]`. $M$ in the notation of the lecture, unless perfect prediction is reached earlier. 

**Note**: 

+ Remember to set the random state for the model in the model's constructor.

```python
from sklearn.ensemble import AdaBoostClassifier

param_grid_ab = ...

model = AdaBoostClassifier(random_state=rng_seed)

result_ab = hypersolve(...)
```

```python
grader.check("q7")
```

# 8. Gradient Boosted Trees

Use the following parameters for the hyper-parameter search. These have the same role as in AdaBoost.
+ `learning_rate: [0.01,0.1]`. 
+ `n_estimators: [20,60,100]`. 

```python
from sklearn.ensemble import GradientBoostingClassifier

param_grid_gbt = ...

model = ...

result_gbt = hypersolve(...)
```

```python
grader.check("q8")
```

# 9. Final model selection

We have now built four separate classifiers: logistic regression, random forest, AdaBoost, and Gradient boosted trees. These are stored below in the `all_models` list. Select from this ist the classifier with the best score. Save the corresponding name, model, and score to variables `best_name`, `best_model`, and `best_score` respectively. Then compute the test accuracy for that model, using its `predict` and `accuracy_score` functions. Store the result as `test_accuracy`.

```python
from sklearn.metrics import accuracy_score

# re-solve logistic regression
param_grid_lr = {
    'model__penalty' : ['l1','l2'],
    'model__C' : np.logspace(-3, 1, 5) }
result_lr = hypersolve( LogisticRegression(solver='liblinear',random_state=rng_seed), param_grid_lr)

all_models = [
    ('logreg',result_lr['best_estimator'], result_lr['best_score']),
    ('rf',result_rf['best_estimator'], result_rf['best_score']),
    ('ab',result_ab['best_estimator'], result_ab['best_score']),
    ('gbt',result_gbt['best_estimator'], result_gbt['best_score']),
]
```

```python
...
best_name = ...
best_model = ...
best_score = ...

ypred = ...
test_accuracy = ...
test_accuracy
```

```python
grader.check("q9")
```

## Submission

Make sure you have run all cells in your notebook in order before running the cell below, so that all images/graphs appear in the output. The cell below will generate a zip file for you to submit. **Please save before exporting!**

```python
# Save your notebook first, then run this cell to export your submission.
grader.export(pdf=False)
```

Tests for q1 (points: None)
```
>>> rawdata.shape == (12528, 9)
True
```
```
>>> np.array(feature_names).shape == (8,)
True
```
```
>>> get_hash(rawdata.values, 2) == '7595d1ad9ea36c042455434471478d67'
True
```
```
>>> [get_hash(a) for a in np.sort(np.array(feature_names))] == ['8a9fb4f90c55e42f5c4d3663d96fa949', '5fb2371942afb4983447b366f8b081d7', '360bdbdde5e6cf2aa449ee1e52fa7c7d', 'c6a62966dd9c0532b2f772a57f57d9df', '9b6b244e08f13ab9da9e39083cb15a63', '63313bd27165b6e021418f49d7730bd6', 'f639989c215ed234d91a1d5c111b4268', '34f484cf1ef5a82db3ad39e45af18642']
True
```


Tests for q2 (points: None)
```
>>> rawdata.shape == (12528, 9)
True
```
```
>>> cleandata.shape == (12528, 6)
True
```
```
>>> [get_hash(a) for a in np.sort(np.array(cleandata.columns))] == ['8a9fb4f90c55e42f5c4d3663d96fa949', '360bdbdde5e6cf2aa449ee1e52fa7c7d', 'c6a62966dd9c0532b2f772a57f57d9df', '63313bd27165b6e021418f49d7730bd6', '34f484cf1ef5a82db3ad39e45af18642', '784fa2495e9e1db533a0d2912a4c8a88']
True
```


Tests for q3 (points: None)
```
>>> bool(N0 > 5000 and N1 > 1000 and (baseline_acc > 0.8))
True
```
```
>>> get_hash(N0, 6) == '58d3d809a09e33275d7d90aab018839d' and get_hash(N1, 6) == 'e8246bf8b9d626bf25e20b1f2c11177e'
True
```
```
>>> get_hash(baseline_acc, 4) == 'efa0f9bcca34237cb56831cca1698fe1'
True
```


Tests for q4 (points: None)
```
>>> Xtrain.shape == (11275, 5) and Xtest.shape == (1253, 5) and (ytrain.shape == (11275,)) and (ytest.shape == (1253,))
True
```
```
>>> get_hash(Xtrain.values.sum(), 5) == '71a15eb580c2f9fc89a457e5cd252852'
True
```
```
>>> get_hash(Xtest.values.sum(), 5) == '336f9c06c9f08ff9a78c3e1c4ec2bd94'
True
```
```
>>> get_hash(ytrain.sum(), 5) == '7404219f5c78f7575defe503bc5f04d2'
True
```
```
>>> get_hash(ytest.sum(), 5) == 'e7878d629de5cac1cabb3ab4b8fd9a42'
True
```


Tests for q5p1 (points: None)
```
>>> {'model', 'scaler'} == logreg_pipe.named_steps.keys()
True
```
```
>>> myparams = logreg_pipe.named_steps['model'].get_params()
>>> myparams['solver'] == 'liblinear' and myparams['random_state'] == rng_seed
True
```
```
>>> [get_hash(z, 4) for z in logreg_pipe.named_steps['scaler'].scale_] == ['8f29de1d80496364e8ad04f2e642631b', '7faa2f235dd84dd374cd179ff36258e2', '5335d2d02b0efd6448ee7c03273c4be1', 'cac91ff20aabe1f6600976bdfeaa9b4a', '75b62c271dc1cb55c54ce2d18f1466d2']
True
```
```
>>> [get_hash(z, 4) for z in logreg_pipe.named_steps['scaler'].mean_] == ['3746bc155c618f8e8604eb220dc6804d', '5661eb581d454be96faf0f02ea5be1a3', 'e5c8e3b00a1278367e806ecac8a07f27', 'ada834294afd71fd7459d48ed83bdadb', 'a2705f8f1cff958e88544f41506035c0']
True
```


Tests for q5p2 (points: None)
```
>>> bp = unpack_gridsearch(gs_logreg)['best_params']
>>> get_hash(bp['model__C'], 4) == '43a1437f7f656cd8be7c996c58719e0a' and get_hash(bp['model__penalty']) == '377fd569971eedeba8fbea28434a390a'
True
```
```
>>> r = unpack_gridsearch(gs_logreg)['scoregrid']
>>> get_hash(r['l1']['mean_test_score'].mean(), 4) == 'd89aeccf02179196f762582380884dc4'
True
```
```
>>> r = unpack_gridsearch(gs_logreg)['scoregrid']
>>> get_hash(r['l2']['mean_test_score'].mean(), 4) == '3d039b3ef68deaa6d2b589ccd81950d0'
True
```


Tests for q6 (points: None)
```
>>> get_hash(result_rf['scoregrid'][3].values.mean(), 4) == '764a15dffce7f2a28e538cd9b9d9dd42'
True
```
```
>>> v = result_rf['scoregrid'][3]['mean_test_score'].values
>>> get_hash(v.mean(), 4) == '764a15dffce7f2a28e538cd9b9d9dd42'
True
```
```
>>> bp = result_rf['best_params']
>>> get_hash(bp['model__max_features']) == '336669dbe720233ed5577ddf81b653d3' and get_hash(bp['model__n_estimators'], 3) == '877583b20242df11218fe61eb534fc09'
True
```


Tests for q7 (points: None)
```
>>> v = result_ab['scoregrid'][0.01]
>>> get_hash(v['mean_test_score'].mean(), 4) == '74901219928841d3a9719f4ba80ab1c5'
True
```
```
>>> v = result_ab['scoregrid'][0.1]
>>> get_hash(v['mean_test_score'].mean(), 4) == 'f35766ef717d880ab29a83831167c3c0'
True
```
```
>>> bp = result_ab['best_params']
>>> get_hash(bp['model__learning_rate'], 3) == 'cb5ae17636e975f9bf71ddf5bc542075' and get_hash(bp['model__n_estimators'], 3) == '56cc33282f4f810d3a644bc652263a7f'
True
```


Tests for q8 (points: None)
```
>>> v = result_gbt['scoregrid'][0.01]
>>> get_hash(v['mean_test_score'].mean(), 4) == 'cc5ac0f4bcd47aa598add0bf923c16a6'
True
```
```
>>> v = result_gbt['scoregrid'][0.1]
>>> get_hash(v['mean_test_score'].mean(), 4) == '09f081b8ee08c11fa7324caa04e16221'
True
```
```
>>> bp = result_gbt['best_params']
>>> get_hash(bp['model__learning_rate'], 3) == 'cb5ae17636e975f9bf71ddf5bc542075' and get_hash(bp['model__n_estimators'], 3) == '75cf3ac5e70c76583be3efb5012bd44e'
True
```


Tests for q9 (points: None)
```
>>> get_hash(best_name) == 'bea2f3fe6ec7414cdf0bf233abba7ef0'
True
```
```
>>> get_hash(test_accuracy, 4) == '39336ee6721cb0aaaaf6fa4bfec2379d'
True
```
