```python
# Initialize Otter
import otter
grader = otter.Notebook("hw5.ipynb")
```

---

<h1><center>SDSE Homework 5<br><br> Time series forecasting </center></h1>

---

In this homework you will apply time series forecasting methods to the problem of predicting demands for a large power grid. The data comes from [PJM](https://www.pjm.com/about-pjm), a regional transmission organization covering a large portion of the east coast of the United States. Specifically, the data is for northern Illinois. It contains hourly values of electricity demand (in MW) between 2004 and 2010. 

The homeworks consists of these parts:

+ Load the data
+ Split the data into training, validation, and testing datasets
+ Normalize the data
+ Organize the data into input sequences and output values
+ Train these models
    + Linear regression
    + Multi-layer perceptron (a.k.a. dense neural network)
    + simple RNN
    + LSTM
+ Compare the models
+ Compute performance of the selected model

```python
# NOTE: If you are running this in JupyterLab, you will have to do the following:
# + Run this cell. This will install tensorflow in your JupyterLab environment.
# + Comment out the code in this cell, to prevent an unnecessary re-install.
# + Restart your kernel (Kernel>Restart kernel...)

%pip install --upgrade pip
%pip install "numpy<1.24"
%pip install --upgrade tensorflow==2.19.0
```

```python
# Do not modify this cell
import os
import random
import numpy as np
import pandas as pd
from pandas.core.indexes.datetimes import DatetimeIndex
import matplotlib.pyplot as plt
import tensorflow
from hashutils import *
os.environ['PYTHONHASHSEED'] = '0'  # optional, for hash-based functions
random.seed(2434)
np.random.seed(2434)
tensorflow.random.set_seed(2434)
tensorflow.keras.utils.set_random_seed(2434)
tensorflow.config.experimental.enable_op_determinism()
os.environ['TF_DETERMINISTIC_OPS'] = '1'
os.environ['TF_CUDNN_DETERMINISTIC'] = '1'
np.set_printoptions(precision=4)
```

### 0. Load the data into a pandas DataFrame

The data is contained in a csv file called `demand.csv`. Load this file into a pandas DataFrame using `read_csv` with the following input arguments:
+ index_col=[0]. The first column in the csv file contains the time stamp. This tells pandas to use this as the index of the data frame.
+ parse_dates=[0]. This tells pandas to convert the time stamps into DateTime objects.

Keep the DataFrame in a variable called `raw_data`.

**Note**: This part has already been completed.

```python
raw_data = pd.read_csv('demand.csv', index_col=[0], parse_dates=[0])

# Display the first 5 rows of the dataset
raw_data.head()
```

### Plot the full time series

```python
plt.figure(figsize=(10,4))
plt.plot(raw_data)
```

### 1. Split the data into training, validation, and testing datasets

Write a function called `split_data` that takes these arguments:
+ `pt`: A number between 0 and 1 corresponding to the proportion of the data to be used for *training*.
+ `pv`: A number between 0 and 1 corresponding to the proportion of the data to be used for *validation*. Validation here means that it will be used by the fitting function of the neural network to evaluate the model at each step of stochastic gradient descent. 

The function should return 3 pandas DataFrames: `data_train`, `data_validate`, and `data_test` (in that order). 

Each of these should contain a **copy** of the segment of `raw_data` that will be used for training, validation, and testing respectively. 

The length of `data_train` should be the integer part of `N*pt`, where `N` is the total number of samples in `raw_data`. Similarly for `data_validate`.

**Notes**:
+ `raw_data` is not passed to `split_data` because it is a global variable.
+ Training data should precede validation data in time. 
+ Validation data should precede testing data in time.

**Hint**:
+ [DataFrame.copy](https://pandas.pydata.org/docs/reference/api/pandas.DataFrame.copy.html)


```python
def split_data(pt,pv):

    assert(pt>=0)
    assert(pv>=0)
    assert(pt+pv<=1)

    ...
    
    return data_train, data_validate, data_test
```

```python
# test your code:
...
```

```python
grader.check("q1")
```

### 2. Split the data

Use your `split_data` function to split the data with the following proportions:
+ `data_train`, with %70 of the data,
+ `data_validate`, with %10 of the data,
+ `data_test`, with %20 of the data.

```python
data_train, data_validate, data_test = ...
```

```python
grader.check("q2")
```

### Plot the training, validation, and test data.

Done already.

```python
plt.figure(figsize=(10,4))
plt.plot(data_train,label='training data')
plt.plot(data_validate,label='validation data')
plt.plot(data_test,label='test data')
plt.legend()
```

### 3. Normalize the data

Use ScikitLearn's `StandardScaler` to normalize the data. Normalizing a dataset, as we've seen, means subtracting its mean and dividing by its standard deviation. This can improve the performance of certain "scale-dependent" models, such as neural networks. (decision trees are scale-*independent*)

Follow these steps:
1. Create a `StandardScaler` object. You can name it whatever you like (e.g. "scaler").
2. Pass the training data (`data_train`) to its `fit` method. This will compute and store the mean and standard deviation of the training sequence. 
3. Apply the scaling transformation to the training sequence by passing it to the scaler's `transform` method. Save the result in a column called `"scaled"` in the training data DataFrame (`data_train`). 
4. Do the same for `data_validate` and `data_test`.

After doing this, `data_train` should look like this:

<img src="df.png" width=350 />

```python
from sklearn.preprocessing import StandardScaler

...
```

```python
grader.check("q3")
```

### Plot the scaled training, validation, and test data. 

Done already.

```python
plt.figure(figsize=(10,4))
plt.plot(data_train['scaled'],label='scaled training data')
plt.plot(data_validate['scaled'],label='scaled validation data')
plt.plot(data_test['scaled'],label='scaled test data')
plt.legend()
```

### 4. Organize the data into input and output sequences

Our goal is to build a model that can forecast the power demand $F$ hours in the future using the $T$ most recent hourly values. To train this model we must organize the training sequence into an input matrix $\mathbf{X}$ and output column vector $\mathbf{Y}$. With $N$ denoting the length of the training sequence (in our case $N=40,915$, as you can verify), the number of rows in $\mathbf{X}$ will be 
$$N' = N-F-T+1$$
The number of columns in $\mathbf{X}$ is $T$.

Write a function called `organize_data` that takes these arguments:
+ `d`: a pandas Series object such as `data_train['scaled']`. 
+ `T`: $T$ as described above.
+ `F`: $F$ as described above.

The function should return the following:
+ `X`, a 2D NumPy array with $N'$ rows and $T$ columns containing the matrix $\mathbf{X}$.
+ `Y`, a 1D NumPy array with $N'$ entries containing the vector $\mathbf{Y}$.
+ `t`, a 1D NumPy array with timestamps for each sample. Each timestamp is a `DatetimeIndex` object.

The two figures below illustrate the organization of $\mathbf{t}$, $\mathbf{X}$, and $\mathbf{Y}$. The first figure shows three input/output samples (green, orange, and blue). The inputs are the sequences of length $T=14$. For example, the green input sequence is collected at 5:00 PM, and it consists of hourly data starting at 4:00 AM. The output of the model is trained to match the forecasted green point at 8:00 PM. 

Similarly, the orange sample consists of an input sequence, collected at 6:00 PM and consisting of  14 hourly values from 5:00 AM to 6:00 PM, and used to predict a value at 9:00 PM.

<img src="samples.png" width=800 />


These samples are arranged into arrays $\mathbf{t}$, $\mathbf{X}$, and $\mathbf{Y}$ as shown below. Notice that the timestamp $\mathbf{t}$ corresponds to the time at which the prediction is made, which is also the time at which the left-most column of $\mathbf{X}$ is measured. For example, the time stamp for the green sample is 5:00 PM.


<img src="snipets.png" width=700 />



```python
def organize_data(d,F,T):

    N = d.shape[0]
    Np = N - F - T + 1
    X = np.empty((Np,T))
    Y = np.empty(Np)
    t = np.empty(Np,dtype=DatetimeIndex)

    # Fill in X, Y, and t
    ...

    return X, Y, t
```

```python
grader.check("q4")
```

### 5. Call `organize_data` on the training, validation, and test data

We wish to predict the power demand $F=12$ hours in the future from the $T=6$ most recent values. Use `organize_data` to assemble $\mathbf{X}$ and $\mathbf{Y}$ arrays for the **scaled** training, validation, and test data. 

Call these:
+ `Xtrain` and `ytrain` for the training data.
+ `Xvalid` and `yvalid` for the validation data.
+ `Xtest` and `ytest` for the test data.

```python
T = ...
F = ...
Xtrain, ytrain, ttrain = ...
Xvalid, yvalid, tvalid = ...
Xtest, ytest, ttest = ...
```

```python
grader.check("q5")
```

### Plot the individual sequences

The plot below shows a few sequences in `Xtrain`, jiggled by a small amount so that they do not overlap.

```python
fig, ax = plt.subplots(figsize=(9,3))
for i in range(50):
    ax.plot(range(i,i+T),Xtrain[i,:]+np.random.normal(scale=0.02),marker='.',linewidth=1)
ax.spines[['top','right','bottom']].set_visible(False)
ax.grid()
```

### 6. Linear regression

Train a Scikit-learn linear regression model using `Xtrain` and `ytrain`. 
+ Do not pass any arguments to the `LinearRegression` constructor. 
+ Save the trained model in the variable `linreg`.


```python
from sklearn.linear_model import LinearRegression

...
```

```python
grader.check("q6")
```

### 7. Performance relative to a baseline model

Let's decide on a performance metric to use. 

Consider the coefficient of determination, $R^2$: 
$$ R^2 = 1 - \frac{\sum(\hat{y}_i-y_i)^2}{\sum(\hat{y}_i-\bar{y})^2}$$
This metric compares the MSE of a model to that of a "baseline" prediction $\bar{y}$, which is simply the mean of the training outputs $\{y_i\}_N$. This makes sense for iid data, but it is too simplistic for time series data. A better baseline for the forecasting problem is to predict that *the future is the same as the present*. In other words, we predict that the value of the power demand in $F$ hours will be the same as the current value. This is illustrated in the figure below for $T=3$ and $F=6$.

<img src="baseline.png" width=700 /> 

The mean absolute error of the model is given by,

$$\text{MAE(model)} = \frac{1}{N'} \sum_{k} \left( |x_{k+F}-\hat{y}_k| \right)    $$

The baseline mean absolute error of the model is,

$$\text{MAE(baseline)} = \frac{1}{N'} \sum_{k} \left( |x_{k+F}-x_k| \right)    $$


The performance of the model is then obtained, analogously to $R^2$ by comparing the MAE of the model to the baseline MAE:

$$ \text{Performance(model)} = 1 - \frac{\text{MAE(model)}}{\text{MAE(baseline)}}$$

Note that the baseline performance is 0, while a "perfect" model (with zero MAE) would have a performance of 1. Better models have larger values of performance.

Write a function called `assess` that implements this performance metric. The function should take these input arguments:
+ `x_current`: a NumPy array with the *current* values of power demand ($x_k$ above). 
+ `y`: a NumPy array with the true future values of power demand ($y_k$ above). 
+ `yhat`: a NumPy array with the model predictions ($\hat{y}_k$ above). 

**Hint**
+ To compute the performance of the linear regression model on the validation data we would run

```python 
assess( x_current = Xtrain[:,-1],
        y = ytrain,
        yhat = linreg.predict(Xtrain))
```

```python
from sklearn.metrics import mean_absolute_error

def assess(x_current,y,yhat):
    ...
    perf = ...
    return perf
```

```python
print(f'Performance of linear regression on the training data: {assess(Xtrain[:,-1],ytrain,linreg.predict(Xtrain)):.4f}')
print(f'Performance of linear regression on the validation data: {assess(Xvalid[:,-1],yvalid,linreg.predict(Xvalid)):.4f}')
```

```python
grader.check("q7")
```

### 8. Include hour-of-day

Our linear regression model is performing only about 30\% closer to "perfect prediction" than the baseline model. Let's see if we can improve it.

 We know that demand for electricity is more or less periodic over a day. Hence, hour-of-day may be a useful input to include in the model. To test this idea, we must first write a function that appends the hour-of-day as a final (right-most) column in our input matrix $\mathbf{X}$. 

Write a function called `append_hour_of_day` that takes these input arguments: 
+ `X`: A 2D NumPy array of input sequences, such as `Xtrain`, `Xvalid`, or `Xtest`.
+ `t`: A 1D NumPy array of time stamps, such as `ttrain`, `tvalid`, or `ttest`.

The function should return a new 2D NumPy array with `X` as its first $T$ columns, and the *hour-of-the-day* in its last column.  The hour-of-the-day is an integer between 0 and 23 inclusive. It is stored as the `hour` attribute of each element in `t`, e.g. `ttrain[0].hour`.

**Hint**:

+ Running `append_hour_of_day(Xtrain,ttrain)` should return an array whose first 5 rows are:

<img src="xtrainhour.png" width=600 /> 



```python
def append_hour_of_day(X,t):
    ...
    return ...
```

```python
# test your code
Xtrain_hour = append_hour_of_day(Xtrain,ttrain)
Xtrain_hour[:5,:]
```

```python
grader.check("q8")
```

### 9. Include hour-of-day in training, validation, and test data

Run your `append_hour_of_day` on the training, validation, and test data. Store the results as 
`Xtrain_hour`, `Xvalid_hour`, and `Xtest_hour` respectively.

```python
...
```

```python
grader.check("q9")
```

### 10. New linear regression with hour-of-day

Train a new linear regression model using the input matrix that includes hour-of-day. Save the fitted model as `linreg_hour`.

Compute the performance of both linear regression models on the validation data. Save the performance of `linreg` to `linreg_perf`, and of `linreg_hour` to `linreg_hour_perf`. Did the performance improve?

```python
...
```

```python
grader.check("q10")
```

### 11. Build an MLP 

We will now try several neural network models to see whether they can deliver better performance than linear regression. We begin with a multi-layer perceptron (MLP).

Build a MLP model for the forecast problem with the hour-of-the-day as an additional input. The model should have this sequence of layers:
1. A dense layer with 64 units and ReLU activation.
2. A dense layer with 32 units and ReLU activation.
3. A dense layer with 16 units and ReLU activation.
4. A dense layer with 1 units no activation (this is the output layer).

Each layer should also be passed a `kernel_initalizer`, which is set to a `GlorotUniform` object. This is done for you in this part, but you should add it to neural networks in parts 13 and 14 as well. This is needed to guarantee repeatability of the results. 

Use variable name `model_mlp` for this model.

Compile the model with its `compile` method and these input arguments:
+ `optimizer="rmsprop"`. Tells Keras to use the "root-mean-squared propagation" variant of stochastic gradient descent for training. 
+ `loss="mse"`: Use the L2 loss function. 
+ `metrics=["mae"]`: Record the mean absolute error at the end of each epoch.

**Note**
+ You will have to pass the `input_shape` to the first layer of the MLP. This should be set to `input_shape=(D,)`, where `D` is the number of features in your training data (i.e. in `Xtrain_hour`) 
+ We also have to pass a "kernel initializer" to each of the layers. This determines the manner in which the weights on each of the edges of the neural network are initialized. We use a random seed to ensure that the initialization is the same every time, so that the autograder can work correctly. 

```python
from tensorflow.keras.models import Sequential
from tensorflow.keras.layers import Dense
from tensorflow.keras.initializers import GlorotUniform


ki = GlorotUniform(seed=2434) # This is needed for reproducibility of the results

model_mlp = Sequential([
    Dense(..., input_shape=..., activation=..., kernel_initializer=ki),
    Dense(..., activation=..., kernel_initializer=ki),
    Dense(..., activation=..., kernel_initializer=ki),
    Dense(..., kernel_initializer=ki)
])

model_mlp.compile( ... )
```

```python
grader.check("q11")
```

### 12. Train the MLP

Train your MLP by calling its `fit` function. Use these input arguments:
+ `x=Xtrain_hour`. Use the training data matrix that includes the hour-of-the-day.
+ `y=ytrain`. Training target.
+ `epochs=10`. Train for 10 epochs.
+ `validation_data=(Xvalid_hour,yvalid)`. Pass in the validation data so that Keras can track the MAE of the model.

Save the output of the fit function to the variable name `history_mlp`.

Use the `assess` function to evaluate the performance of your trained model using the validation data.

**Note**
+ We must copy the code from the previous part into this cell, so that the neural network is created and compiled afresh each time before training. This is to ensure reproducibility of the results and proper function of the autograder. 

```python
random.seed(2434)  # Do not change this. It is needed to ensure repeatability.


### Copy your code from the previous part here ####

###################################################

history_mlp = model_mlp.fit(...)

print('Validation: ',assess(Xvalid[:,-1],yvalid,model_mlp.predict(Xvalid_hour)))
```

```python
grader.check("q12")
```

### Plot the training history for MLP

```python
num_epochs = len(history_mlp.epoch)
plt.figure(figsize=(9,3))
plt.plot(range(num_epochs), history_mlp.history["mae"],marker='o', label="Training MAE")
plt.plot(range(num_epochs), history_mlp.history["val_mae"],marker='o', label="Validation MAE")
plt.xlabel('epoch')
plt.title("Training of MLP")
plt.legend()
```

### 13. Build and train a SimpleRNN

Next, create a neural network with simple RNN units. The model should have this sequence of layers:
1. A SimpleRNN layer with **32** units. Because this is the first layer of the neural netork, we must define the shape of the input. Set `input_shape=(D,1)`. Also, with recurrent layers that feed into downstream recurrent layers, we need to pass `return_sequences=True`.
2. A SimpleRNN layer with **32** units. Again pass `return_sequences=True`.
2. A SimpleRNN layer with **16** units. No need to pass `return_sequences=True` in this case because the next downstream layer is dense.
3. A dense layer with **1** unit and no activation (this is the output layer).

**Notes**:
+ Use variable name `model_srnn` for this model.
+ Do not pass an `activation` to any of the SimpleRNN layers.
+ Compile the model with its `compile` method and the same input arguments as in part 11.
+ Train your simple RNN model by calling its `fit` function. Use the same input arguments as in part 12.
+ Save the output of the fit function to the variable name `history_srnn`.
+ Use the `assess` function to evaluate the performance of your trained model using the validation data. 

```python
from tensorflow.keras.layers import SimpleRNN
random.seed(2434)  # Do not change this. It is needed to ensure repeatability.

model_srnn = ...

print('Validation: ',assess(Xvalid[:,-1],yvalid,model_srnn.predict(Xvalid_hour)))
```

```python
grader.check("q13")
```

### Plot the training history of the Simple RNN

```python
num_epochs = len(history_srnn.epoch)
plt.figure(figsize=(9,3))
plt.plot(range(num_epochs), history_srnn.history["mae"],marker='o', label="Training MAE")
plt.plot(range(num_epochs), history_srnn.history["val_mae"],marker='o', label="Validation MAE")
plt.xlabel('epoch')
plt.title("Training of Simple RNN")
plt.legend()
```

### 14. Build and train an LSTM 

Finally, create a neural network with LSTM units. The instructions are identical to those of the simple RNN.

The model should have this sequence of layers:
1. A LSTM layer with **16** units. Because this is the first layer of the neural netork, we must define the shape of the input. Set `input_shape=(D,1)`. Also, with recurrent layers that feed into downstream recurrent layers, we need to pass `return_sequences=True`.
2. A LSTM layer with **16** units. Again pass `return_sequences=True`.
2. A LSTM layer with **8** units. No need to pass pass `return_sequences=True` in this case because the next downstream layer is dense.
3. A dense layer with **1** unit and no activation (this is the output layer).

**Notes**:
+ Use variable name `model_lstm` for this model.
+ Do not pass an `activation` to any of the LSTM layers.
+ Compile the model with its `compile` method and the same input arguments as in part 11.
+ Train your LSTM model by calling its `fit` function. Use the same input arguments as in part 12.
+ Save the output of the fit function to the variable name `history_lstm`.
+ Use the `assess` function to evaluate the performance of your trained model using the validation data. 



```python
from tensorflow.keras.layers import LSTM
random.seed(2434)  # Do not change this. It is needed to ensure repeatability.

model_lstm = ...
```

```python
grader.check("q14")
```

### Plot the training history of the LSTM

```python
num_epochs = len(history_lstm.epoch)
plt.figure(figsize=(9,3))
plt.plot(range(num_epochs), history_lstm.history["mae"],marker='o', label="Training MAE")
plt.plot(range(num_epochs), history_lstm.history["val_mae"],marker='o', label="Validation MAE")
plt.xlabel('epoch')
plt.title("Training of LSTM")
plt.legend()
```

### Single plot with the training histories of our three neural network models

```python
num_epochs = len(history_lstm.epoch)

plt.figure(figsize=(9,3))
plt.plot(range(num_epochs), history_mlp.history["val_mae"], '.-',markersize=10,label="MLP validation MAE")
plt.plot(range(num_epochs), history_srnn.history["val_mae"], '.-',markersize=10,label="sRNN validation MAE")
plt.plot(range(num_epochs), history_lstm.history["val_mae"], '.-',markersize=10, label="LSTM validation MAE")
plt.title("Validation MAE")
plt.legend()
```

### 15. Pick the best model

Evaluate the performance of the four models (linear regression with hour-of-day, MLP, simple RNN, and LSTM) using the validation data (again). Store the result in these variable names: 
+ `linreg_hour_perf`
+ `mlp_perf`
+ `srnn_perf`
+ `lstm_perf`

Choose the best one and save it as `best_model`. For example, if the linear regression model were the best one, your would write:
```python 
best_model = linreg_hour
```

```python
...
```

```python
grader.check("q15")
```

### 16. Final model test performance

Now that we have made a final selection of a model, it is time to use the test dataset. 

Evaluate the performance of `best_model` on the test data. Save the result to the variable `best_perf`.

```python
...
```

```python
grader.check("q16")
```

## Submission

Make sure you have run all cells in your notebook in order before running the cell below, so that all images/graphs appear in the output. The cell below will generate a zip file for you to submit. **Please save before exporting!**

```python
# Save your notebook first, then run this cell to export your submission.
grader.export(pdf=False)
```

Tests for q1 (points: None)
```
>>> dtrain, dval, dtest = split_data(pt=0.5, pv=0.3)
>>> type(dtrain) == pd.core.frame.DataFrame and type(dval) == pd.core.frame.DataFrame and (type(dtest) == pd.core.frame.DataFrame)
True
```
```
>>> dtrain, dval, dtest = split_data(pt=0.5, pv=0.3)
>>> dtrain.shape == (29225, 1) and dval.shape == (17535, 1) and (dtest.shape == (11690, 1))
True
```
```
>>> dtrain, dval, dtest = split_data(pt=0.5, pv=0.3)
>>> dtrain.index.max() < dval.index.min() and dval.index.max() < dtest.index.min()
True
```
```
>>> dtrain, dval, dtest = split_data(pt=0.7, pv=0.1)
>>> dtrain.iloc[0, 0] = 0
>>> bool(raw_data.iloc[0, 0] == 9198.0)
True
```
```
>>> dtrain, dval, dtest = split_data(pt=0.7, pv=0.1)
>>> dtrain.index.is_monotonic_increasing and dval.index.is_monotonic_increasing and dtest.index.is_monotonic_increasing
True
```


Tests for q10 (points: None)
```
>>> isinstance(linreg_hour, LinearRegression)
True
```
```
>>> hasattr(linreg_hour, 'coef_') and hasattr(linreg_hour, 'intercept_')
True
```
```
>>> get_hash(linreg_hour.coef_, 4) == '041f305f9a93af48e041d3ff6cf628a8' and get_hash(linreg_hour.intercept_, 4) == '05d4938f03ae6ca2613a71cf2cf09bc7'
True
```
```
>>> get_hash(linreg_perf, 4) == 'ff5bb1d297b5bd5a242bf2af6a328b0e' and get_hash(linreg_hour_perf, 4) == 'c982c2eafe1b27be578ba76bdd4b1464'
True
```


Tests for q11 (points: None)
```
>>> l0 = model_mlp.get_config()['layers'][0]
>>> l0['class_name'] == 'InputLayer'
True
```
```
>>> l0 = model_mlp.get_config()['layers'][0]['config']
>>> l0['batch_shape'] == (None, 7)
True
```
```
>>> l1 = model_mlp.get_config()['layers'][1]
>>> l1['class_name'] == 'Dense'
True
```
```
>>> l1 = model_mlp.get_config()['layers'][1]['config']
>>> l1['units'] == 64 and l1['activation'] == 'relu'
True
```
```
>>> l2 = model_mlp.get_config()['layers'][2]
>>> l2['class_name'] == 'Dense'
True
```
```
>>> l2 = model_mlp.get_config()['layers'][2]['config']
>>> l2['units'] == 32 and l2['activation'] == 'relu'
True
```
```
>>> l3 = model_mlp.get_config()['layers'][3]
>>> l3['class_name'] == 'Dense'
True
```
```
>>> l3 = model_mlp.get_config()['layers'][3]['config']
>>> l3['units'] == 16 and l3['activation'] == 'relu'
True
```
```
>>> l4 = model_mlp.get_config()['layers'][4]
>>> l4['class_name'] == 'Dense'
True
```
```
>>> l4 = model_mlp.get_config()['layers'][4]['config']
>>> l4['units'] == 1 and l4['activation'] == 'linear'
True
```
```
>>> cfg = model_mlp.get_compile_config()
>>> 'optimizer' in cfg.keys() and cfg['optimizer']['class_name'] == 'RMSprop'
True
```
```
>>> cfg = model_mlp.get_compile_config()
>>> 'loss' in cfg.keys() and cfg['loss'] == 'mse'
True
```
```
>>> cfg = model_mlp.get_compile_config()
>>> 'metrics' in cfg.keys() and len(cfg['metrics']) == 1 and (cfg['metrics'][0] == 'mae')
True
```


Tests for q12 (points: None)
```
>>> len(history_mlp.epoch) == 10
True
```
```
>>> get_hash(history_mlp.history['mae'], 4) == '180dd7609acd3f40312247e7551658d2'
True
```
```
>>> get_hash(history_mlp.history['val_mae'], 4) == 'a57e2fa3558b8f9e04e7f01b4fc7abc9'
True
```


Tests for q13 (points: None)
```
>>> l0 = model_srnn.get_config()['layers'][0]
>>> l0['class_name'] == 'InputLayer'
True
```
```
>>> l0 = model_srnn.get_config()['layers'][0]['config']
>>> l0['batch_shape'] == (None, 7, 1)
True
```
```
>>> l1 = model_srnn.get_config()['layers'][1]
>>> l1['class_name'] == 'SimpleRNN'
True
```
```
>>> l1 = model_srnn.get_config()['layers'][1]['config']
>>> l1['units'] == 32
True
```
```
>>> l2 = model_srnn.get_config()['layers'][2]
>>> l2['class_name'] == 'SimpleRNN'
True
```
```
>>> l2 = model_srnn.get_config()['layers'][2]['config']
>>> l2['units'] == 32
True
```
```
>>> l3 = model_srnn.get_config()['layers'][3]
>>> l3['class_name'] == 'SimpleRNN'
True
```
```
>>> l3 = model_srnn.get_config()['layers'][3]['config']
>>> l3['units'] == 16
True
```
```
>>> l4 = model_srnn.get_config()['layers'][4]
>>> l4['class_name'] == 'Dense'
True
```
```
>>> l4 = model_srnn.get_config()['layers'][4]['config']
>>> l4['units'] == 1 and l4['activation'] == 'linear'
True
```
```
>>> cfg = model_srnn.get_compile_config()
>>> 'optimizer' in cfg.keys() and cfg['optimizer']['class_name'] == 'RMSprop'
True
```
```
>>> cfg = model_srnn.get_compile_config()
>>> 'loss' in cfg.keys() and cfg['loss'] == 'mse'
True
```
```
>>> cfg = model_srnn.get_compile_config()
>>> 'metrics' in cfg.keys() and len(cfg['metrics']) == 1 and (cfg['metrics'][0] == 'mae')
True
```
```
>>> len(history_srnn.epoch) == 10
True
```
```
>>> get_hash(history_srnn.history['mae'], 4) == 'f5116737e7ed6d38fa5196c037ccdb8f'
True
```
```
>>> get_hash(history_srnn.history['val_mae'], 4) == '4977e75a76d79d0a4d099568c2a36a6a'
True
```


Tests for q14 (points: None)
```
>>> l0 = model_lstm.get_config()['layers'][0]
>>> l0['class_name'] == 'InputLayer'
True
```
```
>>> l0 = model_lstm.get_config()['layers'][0]['config']
>>> l0['batch_shape'] == (None, 7, 1)
True
```
```
>>> l1 = model_lstm.get_config()['layers'][1]
>>> l1['class_name'] == 'LSTM'
True
```
```
>>> l1 = model_lstm.get_config()['layers'][1]['config']
>>> l1['units'] == 16
True
```
```
>>> l2 = model_lstm.get_config()['layers'][2]
>>> l2['class_name'] == 'LSTM'
True
```
```
>>> l2 = model_lstm.get_config()['layers'][2]['config']
>>> l2['units'] == 16
True
```
```
>>> l3 = model_lstm.get_config()['layers'][3]
>>> l3['class_name'] == 'LSTM'
True
```
```
>>> l3 = model_lstm.get_config()['layers'][3]['config']
>>> l3['units'] == 8
True
```
```
>>> l4 = model_lstm.get_config()['layers'][4]
>>> l4['class_name'] == 'Dense'
True
```
```
>>> l4 = model_lstm.get_config()['layers'][4]['config']
>>> l4['units'] == 1 and l4['activation'] == 'linear'
True
```
```
>>> cfg = model_lstm.get_compile_config()
>>> 'optimizer' in cfg.keys() and cfg['optimizer']['class_name'] == 'RMSprop'
True
```
```
>>> cfg = model_lstm.get_compile_config()
>>> 'loss' in cfg.keys() and cfg['loss'] == 'mse'
True
```
```
>>> cfg = model_lstm.get_compile_config()
>>> 'metrics' in cfg.keys() and len(cfg['metrics']) == 1 and (cfg['metrics'][0] == 'mae')
True
```
```
>>> len(history_lstm.epoch) == 10
True
```
```
>>> get_hash(history_lstm.history['mae'], 4) == '8848ccee3942fc58ef6acce8cb93f955'
True
```
```
>>> get_hash(history_lstm.history['val_mae'], 4) == 'd3d9516b9e778bc3a47c45e19524250c'
True
```


Tests for q15 (points: None)
```
>>> bool(np.isclose(mlp_perf, 0.6643458510404325, atol=0.0001))
True
```
```
>>> get_hash(linreg_hour_perf, 4) == 'c982c2eafe1b27be578ba76bdd4b1464'
True
```
```
>>> get_hash(mlp_perf, 4) == '4893e3f17711364b417e77032bc5b1e4' and get_hash(srnn_perf, 4) == '609d69aa26e276cc06ec8b834cc73ab5' and (get_hash(lstm_perf, 4) == '8a52647f06947d3d26632ec454dd25d7')
True
```
```
>>> get_hash(best_model.get_config()['layers'][1]['class_name']) == '9a71a13863c84ba34fc41595fc2ee0c6'
True
```


Tests for q16 (points: None)
```
>>> best_perf > 0.6 and best_perf < 0.7
True
```
```
>>> get_hash(best_perf, 4) == '3fb54c972c18f517326eae8b12ca275a'
True
```


Tests for q2 (points: None)
```
>>> data_train.shape == (40915, 1)
True
```
```
>>> data_validate.shape == (5845, 1)
True
```
```
>>> data_test.shape == (11690, 1)
True
```


Tests for q3 (points: None)
```
>>> set(data_train.columns) == {'power demand [MW]', 'scaled'}
True
```
```
>>> set(data_validate.columns) == {'power demand [MW]', 'scaled'}
True
```
```
>>> set(data_test.columns) == {'power demand [MW]', 'scaled'}
True
```
```
>>> bool(np.isclose(data_train['scaled'].values.mean(), 0.0) and np.isclose(data_train['scaled'].values.std(), 1.0))
True
```
```
>>> bool(np.isclose(data_validate['scaled'].values.mean(), -0.173909) and np.isclose(data_validate['scaled'].values.std(), 0.899079))
True
```
```
>>> bool(np.isclose(data_test['scaled'].values.mean(), -0.07913193) and np.isclose(data_test['scaled'].values.std(), 0.98475571087))
True
```


Tests for q4 (points: None)
```
>>> xa, ya, ta = organize_data(data_train.iloc[:10, 0], F=2, T=2)
>>> xa.shape == (7, 2) and ya.shape == (7,) and (ta.shape == (7,))
True
```
```
>>> xa, ya, ta = organize_data(data_train.iloc[:10, 0], F=2, T=2)
>>> bool(np.isclose(xa.mean(), 8159.785714285715, atol=0.001))
True
```
```
>>> xa, ya, ta = organize_data(data_train.iloc[:10, 0], F=2, T=2)
>>> bool(np.isclose(ya[0], 7917, atol=0.001)) and bool(np.isclose(ya[-1], 9381, atol=0.001))
True
```
```
>>> from pandas import Timestamp
>>> xa, ya, ta = organize_data(data_train.iloc[:10, 0], F=2, T=2)
>>> ta[0] == Timestamp('2004-05-01 03:00:00') and ta[-1] == Timestamp('2004-05-01 09:00:00')
True
```


Tests for q5 (points: None)
```
>>> T == 6 and F == 12
True
```
```
>>> from pandas import Timestamp
>>> isinstance(ttrain[0], Timestamp) and isinstance(tvalid[0], Timestamp) and isinstance(ttest[0], Timestamp)
True
```
```
>>> bool(np.all(np.isclose(Xtrain[0, :], [-1.0773, -1.3391, -1.5005, -1.6114, -1.6486, -1.6577], atol=0.01)))
True
```
```
>>> bool(np.all(np.isclose(Xtrain[-1, :], [0.145, 0.1817, 0.1629, 0.1154, 0.0891, 0.0428], atol=0.01)))
True
```
```
>>> bool(np.all(np.isclose(ytrain[[0, -1]], [-1.0143, -0.7011], atol=0.01)))
True
```
```
>>> Xtrain.shape == (40898, 6) and ytrain.shape == (40898,)
True
```
```
>>> from pandas import Timestamp
>>> ttrain[0] == Timestamp('2004-05-01 07:00:00')
True
```
```
>>> get_hash(Xtrain, 4) == '3013f8ed8a3acc56c34394da4b4f56bb' and get_hash(ytrain, 4) == '9595bff50cf6cdbc82efe15a3caeb857'
True
```
```
>>> get_hash(Xvalid, 4) == 'd8fbf1bc6c5284c7b2d3153c56bc969b' and get_hash(yvalid, 4) == '94beaf31ceb6f129e8fdfcd2b05ea157'
True
```
```
>>> get_hash(Xtest, 4) == 'e8e6916ed45e6bd290710d6273f5181f' and get_hash(ytest, 4) == 'be4b23b8a6d4983b948e4a7a091d0a88'
True
```


Tests for q6 (points: None)
```
>>> isinstance(linreg, LinearRegression)
True
```
```
>>> hasattr(linreg, 'coef_') and hasattr(linreg, 'intercept_')
True
```
```
>>> get_hash(linreg.coef_, 4) == '6688577e040683c9e60f0bc61aaa4f7f'
True
```
```
>>> get_hash(linreg.intercept_, 4) == '15857fea29fe449b2dad03e61cfee3a8'
True
```


Tests for q7 (points: None)
```
>>> bool(np.isclose(assess([1, 2, 3], [5, 5, 6], [7, 8, 9]), 0.2, atol=0.001))
True
```
```
>>> bool(np.isclose(assess(Xvalid[:, -1], yvalid, linreg.predict(Xvalid)), 0.310725112468464, atol=0.001))
True
```
```
>>> bool(np.isclose(assess(Xtest[:, -1], ytest, linreg.predict(Xtest)), 0.3274380065847301, atol=0.001))
True
```


Tests for q8 (points: None)
```
>>> from pandas import Timestamp
>>> bool(np.all(append_hour_of_day(np.ones((3, 2)), [Timestamp('07:00:00'), Timestamp('08:00:00'), Timestamp('09:00:00')]) == np.array([[1.0, 1.0, 7.0], [1.0, 1.0, 8.0], [1.0, 1.0, 9.0]])))
True
```
```
>>> Xtrain_hour = append_hour_of_day(Xtrain, ttrain)
>>> Xvalid_hour = append_hour_of_day(Xvalid, tvalid)
>>> Xtest_hour = append_hour_of_day(Xtest, ttest)
>>> Xtrain_hour.shape == (40898, 7) and Xvalid_hour.shape == (5828, 7) and (Xtest_hour.shape == (11673, 7))
True
```


Tests for q9 (points: None)
```
>>> get_hash(Xtrain_hour, 4) == 'db7ee82f4e09350ecb2cac63a11f9147'
True
```
```
>>> get_hash(Xvalid_hour, 4) == 'fcd6674b42769e0a6230388d8fd8a2f6'
True
```
```
>>> get_hash(Xtest_hour, 4) == '1ecb2a4390e4f39242b7603dc6c7849d'
True
```
